```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04_balance_smallworlds.ipynb")
```

# Homework 4: Network Balance and Small Worlds

```python
!pip install --upgrade networkx
```

```python
from IPython.core.display import HTML
from datascience import *

import matplotlib
matplotlib.use('Agg')
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import os
plt.style.use('fivethirtyeight')

import networkx as nx
from networkx.algorithms import bipartite
```

# Positive and Negative Relationships

Consider the network shown in Figure below: there is an edge between each pair of nodes, with five of the edges corresponding to positive relationships, and the other five of the edges corresponding to negative relationships.
<img src="fig_5_18.png" width="240" height="160" align="center"/>
<br>
Each edge in this network participates in three triangles: one formed by each of the additional nodes who is not already an endpoint of the edge. (For example, the A-B edge participates in a triangle on A, B, and C, a triangle on A, B, and D, and a
triangle on A, B, and E. We can list triangles for the other edges in a similar way.) 

<!-- BEGIN QUESTION -->

## Question 1:
For each edge, how many of the triangles it participates in are balanced, and how many are unbalanced? Please list in the following box.

(Notice that because of the symmetry of the network, the answer will be the same for each positive edge, and also for each negative edge; so it is enough to consider this for one of the positive edges and one of the negative edges.)

**Example:**
1. Edge (AB) participates in three triangles: (AB,BC,AC) is unbalanced, (AB,AD,BD) is unbalanced, (AB,BE,AE) is balanced <br>
2. ... 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 2:
In the social network depicted in the following figure, with each edge labeled as either a strong or weak tie, which nodes satisfy the Strong Triadic Closure Property, and which do not? Provide an explanation for your answer.

<img src="fig_3_22.png" width="360" height="240" align="center"/>
<br>

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Small worlds

In this homework assignment, we're going to explore the concept of *small worlds*. Small worlds have long been studied by social networks researchers, and they have also been discussed in popular culture. As a reminder, the rough idea is that social networks can typically be expected to have two characteristics:

* a high level of clustering
* a short average path length

A high level of clustering is consistent with the idea of triadic closure. And a short average path length is supposed to capture situations we often seem to encounter in our day to day lives: e.g., two strangers find that they have an unexpected acquaintance in common and exclaim "it's a small world!" (see the Milgram article below).

In case you want to read some of the original small world research papers, you can check out some of the papers we talked about in lecture. Here is an article describing an early empirical study by Milgram:

* [Milgram 1967](http://measure.igpp.ucla.edu/GK12-SEE-LA/Lesson_Files_09/Tina_Wey/TW_social_networks_Milgram_1967_small_world_problem.pdf)

And here are a couple of more recent studies in which researchers analyzed mathematical models that can produce networks with small-world properties:

* [Watts & Strogatz 1998](http://www.nature.com/nature/journal/v393/n6684/abs/393440a0.html)
* [Watts 1999](http://www.jstor.org/stable/10.1086/210318?seq=1#page_scan_tab_contents)

First, we'll start by looking at this sample network: 

```python
test_net = nx.Graph([(1,2), (1,3), (2,3), (4,5), (4,6), (4,3), (5,6), (3,5), (2,6), (7,8), (8,9)])
nx.draw_circular(test_net, with_labels=True)
```

<!-- BEGIN QUESTION -->

# Question 3

Edit the table below in a new cell and fill it with the shortest distance between each pair of vertices. 
Enter INF if the nodes are disconnected. 


| &nbsp;  | node 1 | node 2 | node 3 | node 4 | node 5 | node 6 | node 7 | node 8 | node 9 |
|  ------ | -----  | ------ | ------ | ------ | ------ | ------ | ------ | ------ | ------ |
|  node 1 |   -    |  (?)   |   (?)  |    (?) |    (?) |    (?) |    (?) |    (?) |    (?) |
|  node 2 |   -    |  -     |   (?)  |    (?) |    (?) |    (?) |    (?) |    (?) |    (?) |
|  node 3 |   -    |  -     |   -    |    (?) |    (?) |    (?) |    (?) |    (?) |    (?) |
|  node 4 |   -    |  -     |   -    |   -    |    (?) |    (?) |    (?) |    (?) |    (?) |
|  node 5 |   -    |  -     |   -    |   -    |   -    |    (?) |    (?) |    (?) |    (?) |
|  node 6 |   -    |  -     |   -    |   -    |   -    |   -    |    (?) |    (?) |    (?) |
|  node 7 |   -    |  -     |   -    |   -    |   -    |   -    |   -    |    (?) |    (?) |
|  node 8 |   -    |  -     |   -    |   -    |   -    |   -    |   -    |   -    |    (?) |
|  node 9 |   -    |  -     |   -    |   -    |   -    |   -    |   -    |   -    |   -    |  
  

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Question 4 

Now manually calculate the average shortest path length **for each** of the connected components in the network.
  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Question 5 
Now verify the average shortest path length for each of the components in the graph, using the average_shortest_path_length function of the networkx library. For this purpose, one has to iterate over the connected components of the test_net graph.  
  

```python
aspl_vec = []
components = [test_net.subgraph(c) for c in nx.connected_components(test_net)]
for i, each in enumerate(components): 
    aspl = ...
    aspl_vec.append(aspl)
    print ("Component", i, ": Average Shortest Path Length =", aspl )
```

```python
grader.check("q5")
```

<!-- BEGIN QUESTION -->

# Question 6:

True or False: A small world network usually has one giant connected component. Justify and explain your answer.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Please upload the .zip file to Gradescope by Thursday, September 28th at 11:59pm.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q5 (points: 4)
```
>>> (len(aspl_vec) == 2) & (np.all([(x > 1) & (x < 2) for x in aspl_vec]))
True
```
